```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment1.ipynb")
```

## Introduction to Pandas and Scikit-Learn

Pandas is a powerful data manipulation and analysis library for Python. It provides data structures like DataFrames and Series that allow for efficient handling of structured data. Pandas is particularly useful for tasks such as reading and writing data in various formats, data cleaning, merging datasets, and performing complex operations on data.

Scikit-learn, on the other hand, is a machine learning library for Python. It provides a wide range of supervised and unsupervised learning algorithms, as well as tools for model selection, evaluation, and preprocessing. Scikit-learn is designed to be user-friendly and efficient, making it a popular choice for both beginners and experienced data scientists.

Together, Pandas and Scikit-learn form a powerful combination for data analysis and machine learning tasks. Pandas is often used to prepare and manipulate data, which can then be fed into Scikit-learn models for training and prediction.

In this assignment, we'll start with the fundamentals of data loading/manipulation in pandas, then move on to basics of scikit-learn.

## Pandas

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
```

Let's load the **Census Income** dataset which is available on the [UCI ML Dataset Page](https://archive.ics.uci.edu/dataset/20/census+income).

The goal is to predict whether a person's income was greater than $50K based on 1994 census data.

```python
df = pd.read_csv('adult.csv')
```

#### Let's take a quick look at the dataset. We'll use the .head() function to view the first 5 records of the dataset.

```python
df.head()
```

Output:
```
age  workclass  fnlwgt     education  educational-num      marital-status  \
0   25    Private  226802          11th                7       Never-married   
1   38    Private   89814       HS-grad                9  Married-civ-spouse   
2   28  Local-gov  336951    Assoc-acdm               12  Married-civ-spouse   
3   44    Private  160323  Some-college               10  Married-civ-spouse   
4   18          ?  103497  Some-college               10       Never-married   

          occupation relationship   race  gender  capital-gain  capital-loss  \
0  Machine-op-inspct    Own-child  Black    Male             0             0   
1    Farming-fishing      Husband  White    Male             0             0   
2    Protective-serv      Husband  White    Male             0             0   
3  Machine-op-inspct      Husband  Black    Male          7688             0   
4                  ?    Own-child  White  Female             0             0   

   hours-per-week native-country income  
```

## Part 1: Data Loading and pre-processing

### Feel free to refer to the course notes on [pandas](https://tools4ds.github.io/DS701-Course-Notes/02B-Pandas.html) for the following questions.

**Question 1a**: Write a function `get_rows_and_columns` that takes as input a CSV filename, loads this file into a Pandas dataframe, and returns a tuple of the number of rows and columns in the dataframe.

```python
def get_rows_and_columns(file_path):
    df = pd.read_csv('adult.csv')
    return df.shape

```

```python
get_rows_and_columns('adult.csv')
```

Output:
```
(48842, 15)
```

```python
grader.check("q1a")
```

Output:
```
q1a results: All test cases passed!
```

**Question 1b**: Write a function `compute_missing_percentage` that converts '?' to `pd.NA` and returns the percentage of missing data (i.e., NaNs) for each column in the dataset.

The term `pd.NA` is the way to represent [missing values](https://pandas.pydata.org/docs/reference/missing_value.html) (not available) in Pandas.

```python
def compute_missing_percentage(df):
    missingPct = []
    tot = df.shape[0]

    for col in df.columns:
        missingNo = 0

        for i in range(df.shape[0]):
            if str(df.loc[i, col]) == "?":
                df.loc[i, col] = pd.NA
                missingNo += 1
        
        missingPct.append(missingNo / tot)

    output = pd.DataFrame(missingPct, columns=df.columns).T
    return output

```

```python
compute_missing_percentage(df)
```

Output:
```
File ~\AppData\Local\Packages\PythonSoftwareFoundation.Python.3.11_qbz5n2kfra8p0\LocalCache\local-packages\Python311\site-packages\pandas\core\internals\construction.py:845, in to_arrays(data, columns, dtype)
    842     data = [tuple(x) for x in data]
    843     arr = _list_to_arrays(data)
--> 845 content, columns = _finalize_columns_and_data(arr, columns, dtype)
    846 return content, columns

File ~\AppData\Local\Packages\PythonSoftwareFoundation.Python.3.11_qbz5n2kfra8p0\LocalCache\local-packages\Python311\site-packages\pandas\core\internals\construction.py:942, in _finalize_columns_and_data(content, columns, dtype)
    939     columns = _validate_or_indexify_columns(contents, columns)
    940 except AssertionError as err:
    941     # GH#26429 do not raise user-facing AssertionError
--> 942     raise ValueError(err) from err
    944 if len(contents) and contents[0].dtype == np.object_:
    945     contents = convert_object_array(contents, dtype=dtype)

ValueError: 15 columns pa
```

```python
grader.check("q1b")
```

Output:
```
q1b results:
    q1b - 1 result:
        ❌ Test case failed
        Trying:
            bool(np.isclose(compute_missing_percentage(df)['age'], 0.0))
        Expecting:
            True
        **********************************************************************
        Line 1, in q1b 0
        Failed example:
            bool(np.isclose(compute_missing_percentage(df)['age'], 0.0))
        Exception raised:
            Traceback (most recent call last):
              File "C:\Program Files\WindowsApps\PythonSoftwareFoundation.Python.3.11_3.11.2544.0_x64__qbz5n2kfra8p0\Lib\doctest.py", line 1355, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q1b 0[0]>", line 1, in <module>
                bool(np.isclose(compute_missing_percentage(df)['age'], 0.0))
                                ~~~~~~~~~~~~~~~~~~~~~~~~~~~~~~^^^^^^^
            TypeError: 'float' object is not subscriptable

    q1b - 2 result:
        ❌ Test case failed
       
```

**Question 1c**: Write a function `unique_occupation` that returns the number of unique occupation present in the dataset.

```python
def unique_occupation(df):
    ...
```

```python
unique_occupation(df)
```

```python
grader.check("q1c")
```

**Question 1d**: Write a function `categorical_column_with_max_unique_values` that identifies and returns the index of the **column with maximum number of distinct categorical values** in the dataset.

```python
def categorical_column_with_max_unique_values(df):
    ...
```

```python
categorical_column_with_max_unique_values(df)
```

```python
grader.check("q1d")
```

## Part 2: Exploratory data analysis and visualizations

<!-- BEGIN QUESTION -->

**Question 2a**: Write a function `plot_categorical_distribution` to plot the distribution of the column 'education' as a histogram.

You can use Pandas `.plot()` method for this. Look at the DataFrame `.value_counts` method as well. See class examples for how to add labels and titles.

```python
def plot_categorical_distribution(df):
...
```

```python
plot_categorical_distribution(df)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2b**: Write a function `plot_age_hours_scatter` that creates a scatter plot of 'age' vs 'hours-per-week', coloring points by 'income'.

You'll want to look at MatPlotLib's `pyplot.scatter()` for this one.

```python
def plot_age_hours_scatter(df):
...
```

```python
plot_age_hours_scatter(df)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2c**: Write a function `plot_income_by_marital_status` which plots a stacked bar chart that shows the proportion of income levels for each 'marital-status' category.

```python
def plot_income_by_marital_status(df):
    ...
```

```python
plot_income_by_marital_status(df)
```

<!-- END QUESTION -->

## Part 3: Advanced Pandas Operations

**Question 3a**: Write a function `education_stats` that returns a dataframe with mean 'age' and median 'hours-per-week' categorized on the 'education' level.

```python
def education_stats(df):
    ...
```

```python
grader.check("q3a")
```

**Question 3b**: Write a function `calculate_most_popular_occupation` that returns a dataframe of the most popular occupation for each 'native-country' and order them in descending order.

```python
def calculate_most_popular_occupation(df):
    ...
```

```python
calculate_most_popular_occupation(df)
```

```python
grader.check("q3b")
```

**Question 3c**: Write a function `workclass_by_income` that returns a dataframe of the top 5 workclass with the highest number of people having income >50K

```python
def workclass_by_income(df):
     ...
     
```

```python
workclass_by_income(df)
```

```python
grader.check("q3c")
```

## Scikit-learn

```python
# MAKE SURE TO RUN THIS CELL!

import warnings
from sklearn.preprocessing import MinMaxScaler
from sklearn.linear_model import LogisticRegression
from sklearn.model_selection import train_test_split
from sklearn.metrics import mean_squared_error, mean_absolute_error

df_new = df.drop(columns=['native-country', 'fnlwgt']).dropna()
```

In this part, we'll implement [logistic regression](https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LogisticRegression.html) using scikit-learn. Logistic regression is used when trying to predict a binary outcome (0 or 1, True or False, etc.) We will go over the details of logistic regression in details later in the course. 

Here, we will try to predict income (>50k or <= 50k>) and follow standard ML procedures for data pre-processing. You can use scikit-learn's documentation, [the lecture notes on scikit-learn](https://tools4ds.github.io/DS701-Course-Notes/02C-Sklearn.html) or online resources for guidance. 

#### From here on use the 'df_new' variable instead of 'df'

We saw in lecture that models are trained on the 'training' set and evaluated on unseen data on the 'testing' set. The dataset has 'feature' (`X_train`, `X_test`) and the 'outcome' (`y_train`, `y_test`) variables. 

**Question 4a:** Write a function called `split_data` that takes a dataframe as its only parameter, splits it into training and test splits and returns them. Use 20% for the testing set. 

Use `train_test_split` to produce the splits. Provide a `random_state` of 42 for reproducibility.

`split_data` should return 4 things: X_train, X_test y_train and y_test. To do that, you need to pass in the X *and* the y (income column) to `train_test_split`.

```python
def split_data(df):
    ...
```

```python
grader.check("q4a")
```

**Question 4b**: Write a function `preprocess_data` that takes X_train, y_train, X_test, and y_test as input (the splits we created earlier!) and does the following:

 - Scale the *numerical* columns using sklearn's `MinMaxScaler` to the range [0,1] for both train and test sets
  
 - Replace "<=50K" with 0 and ">50K" with 1 in both y_train and y_test

 - One-hot encode the categorical columns for both train and test sets. Check the next cell for some hints! 
  
The function should then return the preprocessed X_train, y_train, X_test, and y_test

Refer to the material below and [sklearn course notes](https://tools4ds.github.io/DS701-Course-Notes/02C-Sklearn.html#prepare-the-dataset) for help! 

One hot encoding is a way of turning textual data into numbers, so that models can work with them. 

Pandas has a method called `pd.get_dummies()` that can do one-hot encoding. 

Let's illustrate with an example.  Let's say we are describing 4 people, each
with attribute 'Gender' and 'City' where they reside.

Create a dictionary with 'City' and 'Gender' keys, each with a length 4 list. 

```python
# Sample DataFrame with categorical columns
data = {'City': ['New York', 'Los Angeles', 'New York', 'Chicago'],
        'Gender': ['Female', 'Male', 'Male', 'Female']}
print(data)
```

We can give that to Pandas `pd.DataFrame()` and it will create a DataFrame with
a City column and a Gender column.

```python
df = pd.DataFrame(data) 
print(df)
```

Now we can call `pd.get_dummies` with the dataframe and it will convert each
categorical column into a set of columns with each category and column entries
of True or False (e.g. 1 or 0).

```python
df_encoded = pd.get_dummies(df, columns=['City', 'Gender'])
# notice how I'm passing in the columns -- you should do this too! Hint: you wrote a function for this earlier. 
```

```python
print(df_encoded)
```

   Notice how now every `City` value has its own column, and that every row with a city has a 1 for that city (row 1 in the old dataframe has New York for the `City`, and row in the new dataframe has a 1 for `City_New York`). Everywhere else you have a 0. 

```python
def preprocess_data(X_train, X_test, y_train, y_test):
    ...
```

```python
grader.check("q4b")
```

Let's get to training! Remember, we're trying to predict whether income is more then 50k (>50k), or less than or equal to (<=50k).

**Question 4c:** Write a function called train_model that takes the training splits (X_train and y_train) as its parameters. 

- Initialize the logistic regression model
- Fit it to our data. (Training step)

At the end, return the fitted model. 

You can refer to [LogisticRegression](https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LogisticRegression.html)

```python
def train_model(X_train, y_train):
    ...
```

```python
grader.check("q4c")
```

Let's evaluate the performance of our model! 

**Question 4d:** Write a function called evaluate_model that takes the fitted model and `X_test`, `y_test` as parameters, runs the model on the testing features (`X_test`) and returns the *accuracy score* of the predictions against the ground truth (`y_test`). 

You can refer to [sklearn metrics](https://scikit-learn.org/stable/modules/model_evaluation.html#regression-metrics).

```python
from sklearn.metrics import accuracy_score

def evaluate_model(model, X_test, y_test):
    ...
```

```python
grader.check("q4d")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Execute all cells and save the notebook before submitting.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1a (points: 2)
```
>>> assert get_rows_and_columns('adult.csv') == (48842, 15)
```


Tests for q1b (points: 2)
```
>>> bool(np.isclose(compute_missing_percentage(df)['age'], 0.0))
True
```
```
>>> bool(np.isclose(compute_missing_percentage(df)['workclass'], 5.730724))
True
```
```
>>> bool(np.isclose(compute_missing_percentage(df)['occupation'], 5.751198))
True
```
```
>>> bool(np.isclose(compute_missing_percentage(df)['native-country'], 1.754637))
True
```
```
>>> int(compute_missing_percentage(df).value_counts().get(0.0)) == 12
True
```


Tests for q1c (points: 2)
```
>>> unique_occupation(df) == 15
True
```
```
>>> unique_occupation(df) == 13
False
```
```
>>> unique_occupation(df) == 14
False
```


Tests for q1d (points: 2)
```
>>> categorical_column_with_max_unique_values(df) == 'native-country'
True
```
```
>>> categorical_column_with_max_unique_values(df) == 'workclass'
False
```
```
>>> categorical_column_with_max_unique_values(df) == 'occupation'
False
```


Tests for q3a (points: 4)
```
>>> len(education_stats(df)) == 16
True
```
```
>>> set(education_stats(df).columns) == {'education', 'age', 'hours-per-week'}
True
```
```
>>> bool(np.isclose(education_stats(df).iloc[0]['age'], 37.902808))
True
```
```
>>> bool(education_stats(df).iloc[1]['hours-per-week'] == 40.0)
True
```
```
>>> int(education_stats(df)['hours-per-week'].value_counts().get(40.0, 0)) == 14
True
```
```
>>> assert education_stats(df).iloc[13]['education'] == 'Preschool'
```
```
>>> assert education_stats(df).iloc[7]['education'] == 'Assoc-acdm'
```


Tests for q3b (points: 3)
```
>>> calculate_most_popular_occupation(df)[calculate_most_popular_occupation(df)['native-country'] == 'United-States']['occupation'].values[0] == 'Exec-managerial'
True
```
```
>>> calculate_most_popular_occupation(df)[calculate_most_popular_occupation(df)['native-country'] == 'India']['occupation'].values[0] == 'Prof-specialty'
True
```
```
>>> calculate_most_popular_occupation(df)[calculate_most_popular_occupation(df)['native-country'] == 'England']['occupation'].values[0] == 'Craft-repair'
False
```
```
>>> calculate_most_popular_occupation(df)[calculate_most_popular_occupation(df)['native-country'] == 'Germany']['occupation'].values[0] == 'Exec-managerial'
False
```


Tests for q3c (points: 3)
```
>>> bool(workclass_by_income(df)[workclass_by_income(df)['workclass'] == 'Private']['count'].values[0] == 7387)
True
```
```
>>> bool(workclass_by_income(df)[workclass_by_income(df)['workclass'] == 'Federal-gov']['count'].values[0] == 561)
True
```
```
>>> bool(workclass_by_income(df)[workclass_by_income(df)['workclass'] == 'Local-gov']['count'].values[0] == 920)
False
```
```
>>> bool(workclass_by_income(df)[workclass_by_income(df)['workclass'] == 'Self-emp-not-inc']['count'].values[0] == 565)
False
```


Tests for q4a (points: 4)
```
>>> def test_split_data():
...     df_new = pd.read_csv('adult.csv').drop(columns=['native-country', 'fnlwgt']).dropna()
...     output = split_data(df_new)
...     assert len(output) == 4, 'split_data() should return 4 things. Make sure you pass both the X and the y to train_test_split.'
...     (X_train, X_test, y_train, y_test) = output
...     assert len(X_train) + len(X_test) == len(df_new), 'The splits do not add up to the original data size.'
...     assert len(y_train) + len(y_test) == len(df_new), 'The splits do not add up to the original data size.'
...     assert np.all(X_train.columns == X_test.columns), 'The columns in X_train and X_test are not the same.'
...     assert len(X_train) > 0 and len(X_test) > 0, 'Some of the splits are empty. Make sure you pass the correct values to train_test_split'
...     assert len(y_train) > 0 and len(y_test) > 0, 'Some of the splits are empty. Make sure you pass the correct values to train_test_split'
>>> test_split_data()
```


Tests for q4b (points: 6)
```
>>> def test_correct_X():
...     with warnings.catch_warnings():
...         warnings.simplefilter('ignore', FutureWarning)
...         (X_train, X_test, y_train, y_test) = split_data(df_new)
...         (X_train, X_test, y_train, y_test) = preprocess_data(X_train, X_test, y_train, y_test)
...     assert np.all(X_train.columns == X_test.columns), 'The columns in X_train and X_test are not the same.'
...     assert np.allclose(np.mean(X_train), np.mean(X_test), atol=0.01), 'The mean of every column is not the same in X_train and X_test.'
...     assert np.all((X_train >= 0) & (X_train <= 1)), 'X_train has values that are not between 0 and 1.'
...     assert np.all((X_test >= 0) & (X_test <= 1)), 'X_test has values that are not between 0 and 1.'
...     assert np.all((np.mean(X_train) >= 0) & (np.mean(X_train) <= 1)), 'X_train has columns with mean values not between 0 and 1.'
...     assert np.all((np.mean(X_test) >= 0) & (np.mean(X_test) <= 1)), 'X_test has columns with mean values not between 0 and 1.'
>>> def test_correct_y():
...     splits = train_test_split(df_new.drop(columns='income'), df_new['income'], test_size=0.2, random_state=42)
...     (X_train, X_test, y_train, y_test) = splits
...     with warnings.catch_warnings():
...         warnings.simplefilter('ignore', FutureWarning)
...         (X_train, X_test, y_train, y_test) = preprocess_data(X_train, X_test, y_train, y_test)
...     assert np.all(y_train.isin([0, 1])), 'y_train has values that are not 0 or 1.'
...     assert np.all(y_test.isin([0, 1])), 'y_test has values that are not 0 or 1.'
>>> test_correct_X()
>>> test_correct_y()
```


Tests for q4c (points: 5)
```
>>> def test_train_model():
...     df_new = pd.read_csv('adult.csv').drop(columns=['native-country', 'fnlwgt']).dropna()
...     (X_train, X_test, y_train, y_test) = split_data(df_new)
...     with warnings.catch_warnings():
...         warnings.simplefilter('ignore')
...         (X_train, X_test, y_train, y_test) = preprocess_data(X_train, X_test, y_train, y_test)
...         model = train_model(X_train, y_train)
...     return isinstance(model, LogisticRegression)
>>> test_train_model()
True
```


Tests for q4d (points: 5)
```
>>> def test_evaluate_model():
...     df_new = pd.read_csv('adult.csv').drop(columns=['native-country', 'fnlwgt']).dropna()
...     (X_train, X_test, y_train, y_test) = split_data(df_new)
...     with warnings.catch_warnings():
...         warnings.simplefilter('ignore')
...         (X_train, X_test, y_train, y_test) = preprocess_data(X_train, X_test, y_train, y_test)
...         model = train_model(X_train, y_train)
...         model = train_model(X_train, y_train)
...         accuracy = evaluate_model(model, X_test, y_test)
...     return bool(np.isclose(accuracy, 0.85, atol=0.1))
>>> test_evaluate_model()
True
```
